A non-conducting role of the Cav1.4 Ca2+ channel drives homeostatic plasticity at the cone photoreceptor synapse

Author response:

The following is the authors’ response to the previous reviews.

Intro.

47-48 rewrite sentence

This sentence has been rewritten as: Photoreceptor synapses are specialized with a vesicle-associated ribbon organelle and postsynaptic neurites of horizontal and bipolar cells that invaginate deep within the terminal

Results

Major comment. Lines 100-103

The new rod data presented here looks like an n = 1. Neither the Results section nor Supp Fig S1, describe the number of cells used. Nor do the authors offer a statistical description with averages, etc.. In addition, the single traces are much improved over their previous study (Maddox et al eLife 2020), but the authors have not described any new approach or trick that improved their rod Ica. Neither Methods section nor Supp section describes the procedure for patching rods (solutions, or Vh which is critical for assessing T-type currents).

Suggestion, if more data exists, then present it. Otherwise, drop the argument.

The recording methodology for recording rods was like that for cones and this has been clarified in the Methods section (lines 725-752). Averaged data (n= at least 5 per group) and statistical analyses have been added to Fig.S1 (renamed Figure 2-Figure Supplement 1), and clearly show that no Ca2+ currents are present in the KI rods.

Supp Fig S2. The legend needs to be fixed. Conversion to PDF file may have created these formatting errors.

This has been corrected (renamed Figure 3-Figure Supplement 2).

Fig 8 a. The position of the light stimulus bar in the KO panel appears to be out of place, shifted too far to the left.

This has been corrected.

Major comments. 219-221

The use of Fluo3-AM is not properly stated here. The text reads "cone pedicles filled with the Ca2+ indicator Fluo3". The wording used could be wrongly interpreted as: whole-cell filling of the cones via patch electrode. However, the Methods section describes bathing the retina in Fluo3-AM, which presumably fills PRs, HCs tips, Mueller glia and bpc dendrites. The Results section should acknowledge that the retina was loaded with Fluo3-AM.

The cell types, and their processes (Muellers, HCs, bpc, PRs), present in a cone pedicle ROI will likely contribute to the Fluo3 readout of Ca2+ in the OPL, because 1) the EM images in Fig 7 highlight how interdigitated the processes are with the presynapse, 2) all express Cav channels, and many if not all express L-Type Cavs in their processes (glia, HC, on-bcs and PRs), and 3) all are depolarized with the addition of high extracellular KCl. The inclusion of Isradipine will inhibit L-type Cavs on pre- and post-synaptic targets, failing to specifically isolate PR Ca2+. Furthermore, Glu Receptor blockers are used here, which would be a great idea if the cones were stimulated with light; however, KCl bypasses the excitatory synaptic pathway and depolarizes all processes within the ROI. Hence, all cellular parts in the ROI will potentially contribute to Fluo3-Ca2+ signals.

Suggestions for presentation of these findings. Ultimately your conclusion is suitable " 233 to 234...... Taken together, our results suggest that Cav3 channels nominally support Ca2+ signals and synaptic transmission in cones of G369i KI mice". The dramatic reduction in Fluo3-Ca2+ signals in the OPL G369i retinas (Fig 9) is a valuable finding for the following reasons: 1) the results do not show a clear compensation from intracellular stores that could potentially supersede the T-type currents in the G369i (which is an argument you make), and 2) there is a massive loss of Ca2+ influx in the OPL of G369i retinas. Since G369i is specific to the PRs, and only cones are present in the mutant G369i, the loss of Fluo3-Ca2+ signal in the mutant ROI reflects in large part loss of cone Fluo3-Ca2+ signals. Your findings illustrate the severity of the mutation, which has also been addressed in the various electro-physio sections of the MS.

Figure 9 also needs to be more clear about 1) the loading of the cells with AM-dye, and 2) the presence of glia, HCs and bc dendrites in the PNA demarcated ROIs.

We regret that we did not make this more clear, but our Fluo 3 loading protocol of whole retina followed by vertical slicing allowed for loading primarily of photoreceptors in the portion of the outer retina that we imaged. We clarified this with the following edit to the text (lines 220-226):

“To test if the diminished HC light responses correlated with lower presynaptic Ca2+ signals in G369i KI cones, we performed 2-photon imaging of vertical slices prepared from whole retina that was incubated with the Ca2+ indicator Fluo3-AM and Alexa-568-conjugated peanut agglutinin (PNA) to demarcate regions of interest (ROIs) corresponding to cone pedicles. With this approach Fluo3 fluorescence was detected only in photoreceptors and ganglion cells and not inner retinal cell-types (e.g., horizontal cells, bipolar cells, Mueller cell soma). Thus, Ca2+ signals reported by Fluo3 fluorescence near PNA-labeling originated primarily from cones.”

We also note that given the considerably larger volume of the cone pedicle relative to the postsynaptic neurites of horizontal and bipolar cells, as well as neighboring glia, it seems unlikely that the latter would contribute significantly to the isradipine-sensitive Ca2+ signal measured in the ROI above the PNA labeling. Moreover, to our knowledge the contribution of Cav1 L-type channels to postsynaptic Ca2+ signals in the dendritic tips of horizontal cells and bipolar cells has not been demonstrated.

Changes from the earlier version to the revised version

Abstract

Abstract

In congenital stationary night blindness type 2 (CSNB2)—a disorder involving the Cav1.4 (L-type) Ca2+ channel—visual impairment is mild considering that Cav1.4 mediates synaptic release from rod and cone photoreceptors. Here, we addressed this conundrum using a Cav1.4 knockout (KO) mouse and a knock-in (G369i KI) mouse expressing a non-conducting Cav1.4. Surprisingly, Cav3 (T-type) Ca2+ currents were detected in cones of G369i KI mice and Cav1.4 KO mice but not in cones of wild-type mouse, ground squirrel, and macaque retina. Whereas Cav1.4 KO mice are blind, G369i KI mice exhibit normal photopic (i.e., cone-mediated) visual behavior. Cone synapses, which fail to form in Cav1.4 KO mice, are present, albeit enlarged, and with some errors in postsynaptic wiring in G369i KI mice. While Cav1.4 KO mice lack evidence of cone synaptic responses, electrophysiological recordings in G369i KI mice revealed nominal transmission from cones to horizontal cells and bipolar cells. In CSNB2, we propose that Cav3 channels maintain cone synaptic output provided that the nonconducting role of Cav1.4 in cone synaptogenesis remains intact. Our findings reveal an unexpected form of homeostatic plasticity that relies on a non-canonical role of an ion channel.

Introduction

Earlier version: At the first synapse in the visual pathway, graded electrical signals produced in rod and cone photoreceptors gate the release of glutamate onto postsynaptic neurons. Photoreceptor synapses are specialized with a ribbon organelle, which helps prime synaptic vesicles, 1, 2 and postsynaptic neurites of horizontal and bipolar cells that invaginate deep within the terminal 3. A variety of proteins interact with the ribbon and synaptic vesicles near release sites (i.e., active zones) 4. The importance of these proteins for vision is illustrated by the numerous inherited retinal diseases linked to mutations in their encoding genes 5.

Revised version: At the first synapse in the visual pathway, graded electrical signals produced in rod and cone photoreceptors gate the release of glutamate onto postsynaptic neurons. Photoreceptor synapses are specialized with a vesicle-associated ribbon organelle and postsynaptic neurites of horizontal and bipolar cells that invaginate deep within the terminal 1. A variety of proteins interact with the ribbon and synaptic vesicles near release sites (i.e., active zones) 2. The importance of these proteins for vision is illustrated by the numerous inherited retinal diseases linked to mutations in their encoding genes 3.

Earlier version: One such gene is CACNA1F, which encodes the voltage-gated Ca2+ (Cav) channel expressed in retinal photoreceptors, Cav1.4 6–8. Among the sub-family of Cav1 L-type channels, Cav1.4 exhibits unusually slow inactivation that is well-matched for supporting the tonic, Ca2+-dependent release of glutamate from photoreceptor synaptic terminals in darkness 9, 10. More than 200 mutations in CACNA1F cause vision disorders including congenital stationary night blindness type 2 (CSNB2) 11, 12. These mutations are broadly categorized as producing a gain of function or loss of function in Cav1.4 13. How these mutations in CACNA1F lead to the variable clinical phenotypes of CSNB2 is largely unknown. Symptoms may include strabismus, low visual acuity, and in many cases, night blindness 14, 15. The latter suggests a primary defect in rod pathways, which is surprising given that Cav1.4 knockout (KO) mice are completely blind and lack any evidence of either rod or cone synaptic responses 6, 7, 16. A major caveat is that rod and cone synapses do not form in Cav1.4 KO mice 7, 16–18. Thus, Cav1.4 KO mice are not suitable for studies of how CACNA1F mutations differentially affect rod and cone pathways or for efforts to uncover how the biophysical properties of Cav1.4 shape photoreceptor synaptic release properties.

Revised version: One such gene is CACNA1F, which encodes the voltage-gated Ca2+ (Cav) channel expressed in retinal photoreceptors, Cav1.4 4–6. Among the sub-family of Cav1 L-type channels, Cav1.4 exhibits unusually slow inactivation that is well-matched for supporting the tonic, Ca2+-dependent release of glutamate from photoreceptor synaptic terminals in darkness 7, 8. More than 200 mutations in CACNA1F cause vision disorders including congenital stationary night blindness type 2 (CSNB2) 9, 10. These mutations are broadly categorized as producing a gain of function or loss of function in Cav1.4 11. How these mutations in CACNA1F lead to the variable clinical phenotypes of CSNB2 is largely unknown. Symptoms may include strabismus, low visual acuity, and in many cases, night blindness 12, 13. The latter suggests a primary defect in rod pathways, which is surprising given that Cav1.4 knockout (KO) mice are completely blind and lack any evidence of either rod or cone synaptic responses 4, 5, 14. A major caveat is that rod and cone synapses do not form in Cav1.4 KO mice 5, 14–16. Thus, Cav1.4 KO mice are not suitable for studies of how CACNA1F mutations differentially affect rod and cone pathways or for efforts to uncover how the biophysical properties of Cav1.4 shape photoreceptor synaptic release properties.

Earlier version: Here, we overcome this hurdle with a knock-in mouse strain (G369i KI) expressing a non-conducting mutant form of Cav1.4 19. We show that, although greatly impaired, cone synapses and downstream signaling through cone pathways can support photopic visual function in G369i KI mice. This novel mechanism requires the ability of the Cav1.4 protein, independent of its Ca2+ conductance, to nucleate the assembly of cone ribbon synapses and involves an aberrant Cav3 (T-type) conductance that appears when Cav1.4 Ca2+ signals are compromised.

Revised version: Here, we overcome this hurdle with a knock-in mouse strain (G369i KI) expressing a non-conducting mutant form of Cav1.4 17. We show that, although greatly impaired, cone synapses and downstream signaling through cone pathways can support photopic visual function in G369i KI mice. This novel mechanism requires the ability of the Cav1.4 protein, independent of its Ca2+ conductance, to nucleate the assembly of cone ribbon synapses and involves an aberrant Cav3 (T-type) conductance that appears when Cav1.4 Ca2+ signals are compromised.

Results

Cone pedicles and ribbons are present in G369i KI mice but not in Cav1.4 KO mice

Earlier version: The G369i mutation results in an insertion of a glycine residue in a transmembrane domain, which prevents Ca2+ permeation through Cav1.419. The presence of rod synapses in G369i KI mice but not Cav1.4 KO mice indicates that the Cav1.4 protein and not its Ca2+ conductance is required for rod synapse assembly19. To test if this is also true for cones, we performed immunofluorescence and confocal analyses using antibodies against cone arrestin (CAR) and CtBP2 to label cone terminals (i.e., pedicles) and ribbons, respectively (Fig. 1). In Cav1.4 KO mouse retinas, cone pedicles were shrunken and retracted into the outer nuclear layer (ONL, Fig. 1a) with little evidence of elongated ribbons (Fig. 1b). In contrast, cone pedicles in G369i KI mice were normally localized in the outer plexiform layer (OPL, Fig. 1a) and were populated by multiple ribbons that were associated with labeling for Cav1.4 (Fig. 1b). Some abnormalities were apparent in G369i KI cone pedicles, such as extremely elongated ribbons and telodendria extending towards the inner nuclear layer (Fig.1b). These results indicate that Cav1.4 Ca2+ signals are dispensable for the integrity of ribbons and the pedicle but are necessary for their overall refinement.

Revised version: The G369i mutation results in an insertion of a glycine residue in a transmembrane domain, which prevents Ca2+ permeation through Cav1.417. The presence of rod synapses in G369i KI mice but not Cav1.4 KO mice indicates that the Cav1.4 protein and not its Ca2+ conductance is required for rod synapse assembly17. To test if this is also true for cones, we performed immunofluorescence and confocal analyses using antibodies against cone arrestin (CAR) and CtBP2 to label cone terminals (i.e., pedicles) and ribbons, respectively (Fig. 1). In Cav1.4 KO mouse retinas, cone pedicles were shrunken and retracted into the outer nuclear layer (ONL, Fig. 1a) with little evidence of elongated ribbons (Fig. 1b). In contrast, cone pedicles in G369i KI mice were normally localized in the outer plexiform layer (OPL, Fig. 1a) and were populated by multiple ribbons that were associated with labeling for Cav1.4 (Fig. 1b). Some abnormalities were apparent in G369i KI cone pedicles, such as extremely elongated ribbons and telodendria extending towards the inner nuclear layer (Fig.1b). These results indicate that Cav1.4 Ca2+ signals are dispensable for the integrity of ribbons and the pedicle but are necessary for their overall refinement.

Cone pedicles and the mutant Cav1.4 G369i channel are normally localized in the retina of G369i KI mice but not Cav1.4 KO mice.

Confocal images of the outer nuclear layer (ONL) and outer plexiform layer (OPL) of retina from WT, G369i KI, and Cav1.4 KO mice that were labeled with antibodies against cone arrestin (CAR), CtBP2, and Cav1.4. a, Inverted images of CAR labeling. Lower panels depict pedicles labeled by CAR antibodies (dotted outlines) and correspond to the boxed regions in the upper panels. Cone pedicles remain in the OPL of WT and G369i KI retina (arrows) but are misshapen and retracted into the ONL of the CaV1.4 KO retina (open arrowheads). Solid arrowheads indicate telodendria which extend only apically in WT retina but extend basally and laterally in G369i KI retina. b, Deconvolved confocal images showing CaV1.4 labeling near cone ribbons in WT and G369i KI pedicles (arrows) and ribbon spheres without CaV1.4 labeling in the CaV1.4 KO pedicle (arrowheads). Rod-associated CtBP2 and Cav1.4 labeling was removed for clarity.

Cav3 channel activity is present in cones of G369i KI and Cav1.4 KO mice but not WT mice, ground squirrel, or macaque retina

Earlier version: A prevailing yet unsupported hypothesis regarding the relatively mild visual phenotypes in CSNB2 is that additional Cav subtypes may compensate for Cav1.4 loss-of-function in cones. If so, then Ca2+ currents (ICa) mediated by these subtypes should be evident in cones of Cav1.4 KO and G369i KI mice. To test this, we performed patch clamp recordings of cones in retinal slices of adult WT, G369i KI, and Cav1.4 KO mice under conditions designed to isolate ICa. During voltage ramps in WT cones, ICa activated around −50 mV and peaked near −20 mV, consistent with the properties of Cav1.4 (Fig. 2a-d). While not observed in rods of G369i KI or Cav1.4 KO mice19, a small-amplitude ICa that activated around −60 mV and peaked near −35 mV was detected in cones of these mice (Fig.2a-d). To uncover this low voltage-activated ICa, we used a faster voltage ramp (0.5 mV/ms) than in the previous study of G369i KI rods (∼0.15 mV/ms)19. With this faster voltage ramp, ICa was not apparent in G369i KI rods (Supp.Fig.S1). Therefore, the aberrant ICa is an adaptation to Cav1.4 loss-of-function in cones but not rods.

Revised version: A prevailing yet unsupported hypothesis regarding the relatively mild visual phenotypes in CSNB2 is that additional Cav subtypes may compensate for Cav1.4 loss-of-function in cones. If so, then Ca2+ currents (ICa) mediated by these subtypes should be evident in cones of Cav1.4 KO and G369i KI mice. To test this, we performed patch clamp recordings of cones in retinal slices of adult WT, G369i KI, and Cav1.4 KO mice under conditions designed to isolate ICa. During voltage ramps in WT cones, ICa activated around −50 mV and peaked near −20 mV, consistent with the properties of Cav1.4 (Fig. 2a-d). While not observed in rods of G369i KI or Cav1.4 KO mice17, a small-amplitude ICa that activated around −60 mV and peaked near −35 mV was detected in cones of these mice (Fig.2a-d). To uncover this low voltage-activated ICa, we used a faster voltage ramp (0.5 mV/ms) than in the previous study of G369i KI rods (∼0.15 mV/ms)17. With this faster voltage ramp, ICa was not apparent in G369i KI rods (Fig.2 Fig.Supp. 1). Therefore, the aberrant ICa is an adaptation to Cav1.4 loss-of-function in cones but not rods.

A low-voltage activated ICa is present in cones of G369i KI and CaV1.4 KO but absent in cones of WT mice.

Earlier version: a, Representative traces of ICa evoked by voltage ramps. b, ICa traces from a normalized to their peak current amplitude to illustrate the hyperpolarizing shift in the I-V from cones of G369i KI and CaV1.4 KO mice. c,d, I-V (c) and G-V (d) relationships for ICa evoked by 50 ms, +5 mV increments from a holding voltage of −90 mV. Vm, test voltage. WT, n = 13; G369i KI, n = 9; CaV1.4 KO, n = 3. e, Representative ICa traces and voltage protocol (bottom) for steady-state inactivation. ICa was evoked by a conditioning pre-pulse from −90 mV to various voltages for 500 ms followed by a test pulse to −30 mV for 50 ms. f, Steady-state inactivation data from cones as recorded in e. I/Imax represents the current amplitude of each −30-mV test pulse (I) normalized to current amplitude of the −30-mV test pulse preceded by a −90 mV conditioning pre-pulse (Imax) and was plotted against pre-pulse voltage. Shaded region indicates window current for G369i KI cones. Line without symbols represents G-V curve for G369i KI cones replotted from d. WT, n = 8; G369i KI, n = 8; Cav1.4 KO, n = 3. In graphs c,d, and f, smooth lines represent Boltzmann fits, symbols and bars represent mean ± SEM, respectively. g, Representative ICa traces from WT and G369i KI cones and voltage protocol (top). ICa was evoked by a step from - 100 mV to −30 mV before (P1) or after (P2) a 500 ms step to −60 mV. h,i, Graphs depicting peak amplitude of ICa during P1 and P2 steps (h) and the ratio of the amplitudes of ICa evoked by P2 and P1 pulses (P2/P1; i) from cones recorded as in g. WT, n = 6; G369i KI, n = 6, *p < 0.05 by paired t-test.

Revised version: a, Representative traces of ICaevoked by voltage ramps. b, ICa traces from a normalized to their peak current amplitude to illustrate the hyperpolarizing shift in the I-V from cones of G369i KI and CaV1.4 KO mice. c,d, I-V (c) and G-V (d) relationships for ICaevoked by 50 ms, +5 mV increments from a holding voltage of −90 mV. Vm, test voltage. WT, n = 13; G369i KI, n = 9; CaV1.4 KO, n = 3. e, Representative ICa traces and voltage protocol (bottom) for steady-state inactivation. ICa was evoked by a conditioning pre-pulse from −90 mV to various voltages for 500 ms followed by a test pulse to −30 mV for 50 ms. f, Steady-state inactivation data from cones as recorded in e. I/Imaxrepresents the current amplitude of each - 30-mV test pulse (I) normalized to current amplitude of the −30-mV test pulse preceded by a −90 mV conditioning pre-pulse (Imax) and was plotted against pre-pulse voltage. Shaded region indicates window current for G369i KI cones. Line without symbols represents G-V curve for G369i KI cones replotted from d. WT, n = 8; G369i KI, n = 8; Cav1.4 KO, n = 3. In graphs c,d, and f, smooth lines represent Boltzmann fits, symbols and bars represent mean ± SEM, respectively. g, Representative ICa traces from WT and G369i KI cones and voltage protocol (top). ICa was evoked by a step from −100 mV to −30 mV before (P1) or after (P2) a 500 ms step to −60 mV. h,i, Graphs depicting peak amplitude of ICa during P1 and P2 steps (h) and the ratio of the amplitudes of ICa evoked by P2 and P1 pulses (P2/P1; i) from cones recorded as in g. WT, n = 6; G369i KI, n = 6, *p < 0.05 by paired t-test.

Earlier version: We next used step depolarizations to identify the type of Cav channel underlying ICa in G369i KI and Cav1.4 KO cones. Analysis of current-voltage (I-V) and conductance-voltage (G-V) relationships revealed a significant, hyperpolarizing shift (∼10 mV) in the voltage of half maximal activation (Vh) in G369i KI cones (Fig.2c,d; Table 1). ICa in G369i KI cones was not sustained as in WT cones but inactivated rapidly during 500-ms step depolarizations (Fig. 2e). For G369i KI cones, overlay of the conductance-voltage (Fig. 2d) and inactivation curves revealed a sizeable window current (Fig. 2f). These features of ICa in G369i KI and Cav1.4 KO cones resembled those of Cav3 T-type channels rather than Cav1.420.

Revised version: We next used step depolarizations to identify the type of Cav channel underlying ICa in G369i KI and Cav1.4 KO cones. Analysis of current-voltage (I-V) and conductance-voltage (G-V) relationships revealed a significant, hyperpolarizing shift (∼10 mV) in the voltage of half maximal activation (Vh) in G369i KI cones (Fig.2c,d; Table 1). ICa in G369i KI cones was not sustained as in WT cones but inactivated rapidly during 500-ms step depolarizations (Fig. 2e). For G369i KI cones, overlay of the conductance-voltage (Fig. 2d) and inactivation curves revealed a sizeable window current (Fig. 2f). These features of ICain G369i KI and Cav1.4 KO cones resembled those of Cav3 T-type channels rather than Cav1.418.

Earlier version: In a previous study, a Cav3-like current was detected in patch clamp recordings of cone pedicles in WT mouse cones21. However, in recordings of WT mouse cone somas, we did not observe a low-voltage activated component in the I-V or G-V curves that would be indicative of a Cav3 subtype (Fig.2c,d). To further test for a Cav3 contribution to ICa in WT mouse cones, we used a double pulse voltage protocol where ICa was evoked before (P1) and after (P2) a 500-ms step to −60 mV (Fig.2g). With this protocol, the inactivation of Cav3 channels by the −60 mV step should lead to a reduction in the amplitude of the P2 vs the P1 current. While such a reduction was clearly evident in cones of G369i KI mice (i.e., P2/P1 < 1), there was no significant difference in the P2 and P1 currents in cones of WT mice (Fig.2g-i).

Revised version: In a previous study, a Cav3-like current was detected in patch clamp recordings of cone pedicles in WT mouse cones19. However, in recordings of WT mouse cone somas, we did not observe a low-voltage activated component in the I-V or G-V curves that would be indicative of a Cav3 subtype (Fig.2c,d). To further test for a Cav3 contribution to ICa in WT mouse cones, we used a double pulse voltage protocol where ICa was evoked before (P1) and after (P2) a 500-ms step to −60 mV (Fig.2g). With this protocol, the inactivation of Cav3 channels by the −60 mV step should lead to a reduction in the amplitude of the P2 vs the P1 current. While such a reduction was clearly evident in cones of G369i KI mice (i.e., P2/P1 < 1), there was no significant difference in the P2 and P1 currents in cones of WT mice (Fig.2g-i).

Earlier version: We also used antagonists selective for Cav1 and Cav3 channels (isradipine and ML 218, respectively). As expected, ICa in WT cones was significantly but incompletely suppressed by isradipine (Fig. 3a,b). The Vh of the I-V was similar before (−36.5 ± 0.8 mV) and after (−38.3 ± 1.8 mV) perfusion of isradipine (t(3,4)=0.90, p = 0.43 by paired t-test), suggesting that the residual ICa after exposure to isradipine arose from unblocked Cav1.4 rather than Cav3 channels. While ML 218 modestly inhibited ICa in some WT cones, the effect was not statistically significant (Fig. 3c,d) and could be attributed to weak activity on Cav1.4. In transfected HEK293T cells, ML 218 modestly inhibited ICa and caused a negative shift in the Vh for Cav1.4, in contrast to its strong inhibition of Cav3.2 (Supp.Fig. S2). Like its effects on Cav3.2, ML 218 nearly abolished ICa in cones of G369i KI mice whereas isradipine had little effect (Fig. 3a-d).

Revised version: We also used antagonists selective for Cav1 and Cav3 channels (isradipine and ML 218, respectively). As expected, ICa in WT cones was significantly but incompletely suppressed by isradipine (Fig. 3a,b). The Vh of the I-V was similar before (−36.5 ± 0.8 mV) and after (−38.3 ± 1.8 mV) perfusion of isradipine (t(3,4)=0.90, p = 0.43 by paired t-test), suggesting that the residual ICa after exposure to isradipine arose from unblocked Cav1.4 rather than Cav3 channels. While ML 218 modestly inhibited ICa in some WT cones, the effect was not statistically significant (Fig. 3c,d) and could be attributed to weak activity on Cav1.4. In transfected HEK293T cells, ML 218 modestly inhibited ICa and caused a negative shift in the Vhfor Cav1.4, in contrast to its strong inhibition of Cav3.2 (Fig.3-Fig.Supp.2). Like its effects on Cav3.2, ML 218 nearly abolished ICain cones of G369i KI mice whereas isradipine had little effect (Fig. 3a-d).

Pharmacological characterization of ICa in cones of WT and G369i KI mice.

Earlier version: a, Representative traces for ICa evoked by voltage ramps in cones of retinas from WT (top) and G369i KI (bottom) mice before (baseline) and after exposure to isradipine (ISR, 1 μM). Middle panel, WT ICa from voltage ramps in the top panel were normalized to their peak ICa to clarify the similar properties of ICa before and after ISR exposure. b, Peak ICa before (−) and during (+) ISR exposure from cones as recorded in a. WT, n = 4; G369i KI, n = 5; *p < 0.05 by paired t-test. c, Representative traces for ICa evoked by voltage ramps in cones of retinas from WT or G369i KI mice before (baseline) and after exposure to ML 218 (5 μM). Middle panel, WT ICa from voltage ramps in the top panel were normalized to their peak ICa to clarify the similar properties of ICa before and after ML 218 exposure. d,

Revised version: a, Representative traces for ICa evoked by voltage ramps in cones of retinas from WT (top) and G369i KI (bottom) mice before (baseline) and after exposure to isradipine (ISR, 1 μM). Middle panel, WT ICa from voltage ramps in the top panel were normalized to their peak ICa to clarify the similar properties of ICa before and after ISR exposure. b, Peak ICa before (−) and during (+) ISR exposure from cones as recorded in a. WT, n = 4; G369i KI, n = 5; *p < 0.05 by paired t-test. c, Representative traces for ICaevoked by voltage ramps in cones of retinas from WT or G369i KI mice before (baseline) and after exposure to ML 218 (5 μM). Middle panel, WT ICa from voltage ramps in the top panel were normalized to their peak ICa to clarify the similar properties of ICa before and after ML 218 exposure. d, Peak ICa before (−) and during (+) ML 218 exposure. WT, n = 5; G369i KI, n = 5; *p < 0.05 by paired t-test.

Removed in the revised version: Peak ICa before (−) and during (+) ML 218 exposure. WT, n = 5; G369i KI, n = 5; *p < 0.05 by paired t-test.

Earlier version: Finally, we recorded from ground squirrel and macaque cones, where the large amplitude of ICa facilitates pharmacological and biophysical analyses. Consistent with its actions on Cav1.4 in WT mouse cones (Fig. 3b) and transfected HEK293T cells (Supp.Fig. S2), ML 218 caused an insignificant inhibition of peak ICa (−7.0 ± 20.8%, n = 6 cones) as well as a negative shift in the voltage-dependence of activation in ground squirrel cones (ΔV1/2 = −1.16 ± 0.94 mV, n = 6, mean ± SD, p = 0.029, t-test; Fig. 4a,b). By contrast, application of isradipine dramatically suppressed ICa in cones in various regions of the ground squirrel retina (Fig. 4c,d). Subsequent application of ML 218 did not reveal a Cav3-like current: the I-V and G-V of the isradipine + ML 218-sensitive current were shifted in the positive rather than the negative direction relative to the isradipine-sensitive current (Fig.4c-g). This result is consistent with the time- and voltage-dependent block of Cav1 channels by dihydropyridine antagonists such as isradipine 22. As a positive control, we confirmed that ML 218 blocked a prominent Cav3-type current in ground squirrel type 3a OFF cone bipolar cells (Fig.4h,i). In recordings of macaque cone pedicles, we compared the properties of ICa using a holding voltage of −90 mV or −50 mV, the latter of which should inactivate Cav3 channels (Fig.5a). There was no difference in the I-V or G-V relationships obtained at these holding voltages (Fig.5b-e). As in WT mouse cones (Fig.2g-i), there was also no difference in ICa recorded before and after an inactivating pulse to −60 mV (Fig.5f,g). We conclude that Cav3 channels do not normally contribute to ICa in cones of WT mice, ground squirrel, and macaque.

Revised version: Finally, we recorded from ground squirrel and macaque cones, where the large amplitude of ICa facilitates pharmacological and biophysical analyses. Consistent with its actions on Cav1.4 in WT mouse cones (Fig. 3b) and transfected HEK293T cells (Fig.3-Fig.Supp.2), ML 218 caused an insignificant inhibition of peak ICa (−7.0 ± 20.8%, n = 6 cones) as well as a negative shift in the voltage-dependence of activation in ground squirrel cones (ΔV1/2 = −1.16 ± 0.94 mV, n = 6, mean ± SD, p = 0.029, t-test; Fig. 4a,b). By contrast, application of isradipine dramatically suppressed ICain cones in various regions of the ground squirrel retina (Fig. 4c,d). Subsequent application of ML 218 did not reveal a Cav3-like current: the I-V and G-V of the isradipine + ML 218-sensitive current were shifted in the positive rather than the negative direction relative to the isradipine-sensitive current (Fig.4c-g). This result is consistent with the time- and voltage-dependent block of Cav1 channels by dihydropyridine antagonists such as isradipine 20. As a positive control, we confirmed that ML 218 blocked a prominent Cav3-type current in ground squirrel type 3a OFF cone bipolar cells (Fig.4h,i). In recordings of macaque cone pedicles, we compared the properties of ICausing a holding voltage of −90 mV or −50 mV, the latter of which should inactivate Cav3 channels (Fig.5a). There was no difference in the I-V or G-V relationships obtained at these holding voltages (Fig.5b-e). As in WT mouse cones (Fig.2g-i), there was also no difference in ICa recorded before and after an inactivating pulse to - 60 mV (Fig.5f,g). We conclude that Cav3 channels do not normally contribute to ICa in cones of WT mice, ground squirrel, and macaque.

Pharmacological characterization of ICa in ground squirrel cones.

a, Representative traces corresponding to baseline-corrected ICa evoked by voltage ramps before (baseline) and during application of ML 218. b, G-V relationship of ICa in a. c, Representative traces corresponding to ICa evoked by voltage ramps in control (baseline) and in ISR (2 μM) alone or in ISR+ML 218 (5 μM). d, Peak ICa from cones including the record from c before (−) and during (+) ISR or ISR+ML 218 block. The changes during the addition of ML218 were small but significant (sup., n = 7, ISR: p < 0.0001, ML218: p = 0.0047; mid., n = 5, ISR: p < 0.0001, ML218: p = 0.0001; inf., n = 5, ISR: p = 0.0019, ML218: p = 0.0460; two-tailed t test) and were likely due to a further non-specific, time-dependent reduction in the Cav1 current. e, ICa was evoked by steps from −85 mV to voltages between −65 and −5 mV in increments of +10 mV (protocol shown below the current traces in f). Representative traces are shown in control, ISR, or ISR+ML 218. Current traces during steps to −55 and −45 mV lacked both transient and ML 218-senstive components. Dashed lines indicate zero current. f, Traces show the ICa sensitive to ISR (top) and ISR+ML218 (middle). The ICa recorded in ISR was subtracted from the baseline ICa (top). ICa recorded in ISR+ML218 was subtracted from ICa recorded in ISR (middle). g, Top, I-V relationship of peak and steady-state ICa from f. Bottom, G-V relationship of the ICa blocked by ISR and ISR + ML218. Smooth lines represent Boltzmann fits. Symbols and bars represent mean ± SEM (n = 4 cones). Due a negative shift in the activation properties caused by ML 218 on a residual Cav1 current that is assumed to remain at the end of the experiment, the G-V curve of the ICa isolated through subtraction by applying ISR+ML 281 underwent a statistically insignificant shift to the right. Data presented in e-g are from the same cone. h, Representative ICa traces from voltage ramps in an OFF cb3a bipolar cell before and during the application of ML 218. V1/2 was shifted to the right by 13.1 ± 4.3 mV (n = 3 cb3a cells, mean ± SD) consistent with the block of a Cav3-like conductance. In cb3b OFF bipolar cells, ML 218 produced only a slight leftward shift in an ICa that had a more depolarized activation range, consistent with the exclusive expression a Cav1-type current (ΔV1/2 = −1.9 ± 0.4 mV, n = 3 cells; data not shown). i, Neurobiotin fill shows cb3a axon stratification and morphology in retina labeled with antibodies against bassoon and choline acetyltransferase (CHAT) to label the OFF and ON sublamina of the IPL.

Characterization of ICa in patch clamp recordings of cone pedicles of macaque retina.

a-c, Representative current traces (a) and corresponding I-V (b) and G-V (c) relationships for ICa evoked by 200-ms steps from a holding voltage of −90 mV or −50 mV. n = 3 cones. d,e, V1/2 (d) and slope factor (e) obtained from Boltzmann fit of data in c. p = 0.50 in d; p = 0.25 in e by Wilcoxon matched pairs signed rank test. f, Voltage protocol (top) and representative ICa traces (bottom left) from cones as recorded in Fig. 2g-i. Bottom right, expanded view of the boxed region in the left traces. g, Peak ICa during P1 and P2 steps from cones recorded as in f. n = 5 cones, p = 0.62 by Wilcoxon matched pairs signed rank test.

Earlier version: To test whether the appearance of the Cav3-mediated ICa in G369i KI cones could result from an increase in the expression of a specific Cav3 subtype (i.e., Cav3.1, Cav3.2, or Cav3.3), we initially tried labeling with commercially available antibodies against Cav3.2. However, the data were unreliable since these antibodies produced a similar labeling pattern in Cav3.2 KO retinal tissue (data not shown). As an alternative, we performed single cell RNA sequencing on cells isolated from the retina of WT and G369i KI mice (Supp. Fig.S3). We sequenced 28,782 cells and, with unsupervised clustering at resolution=0.5, we grouped the cells into 20 clusters. Clusters 8 and 16 were identified as cones based on the expression of cone marker genes (e.g., Opn1sw, Opn1mw, Arr3, Gnat2). Consistent with previous studies21, 23, Cacna1h encoding the Cav3.2 subtype was the major Cav3 transcript expressed in WT mouse cones. There was no significant difference in Cacna1h expression between WT and G369i KI cones according to this analysis (Supp. Fig.S3). Thus, the enhanced activity of Cav3 channels in cones of G369i KI mice results from a mechanism other than increased transcription of Cacna1h.

Revised version: To test whether the appearance of the Cav3-mediated ICain G369i KI cones could result from an increase in the expression of a specific Cav3 subtype (i.e., Cav3.1, Cav3.2, or Cav3.3), we initially tried labeling with commercially available antibodies against Cav3.2. However, the data were unreliable since these antibodies produced a similar labeling pattern in Cav3.2 KO retinal tissue (data not shown). As an alternative, we performed single cell RNA sequencing on cells isolated from the retina of WT and G369i KI mice (Fig.2-Fig.Supp.3). We sequenced 28,782 cells and, with unsupervised clustering at resolution=0.5, we grouped the cells into 20 clusters. Clusters 8 and 16 were identified as cones based on the expression of cone marker genes (e.g., Opn1sw, Opn1mw, Arr3, Gnat2). Consistent with previous studies19, 21, Cacna1h encoding the Cav3.2 subtype was the major Cav3 transcript expressed in WT mouse cones. There was no significant difference in Cacna1h expression between WT and G369i KI cones according to this analysis (Fig.2-Fig.Supp.3). Thus, the enhanced activity of Cav3 channels in cones of G369i KI mice results from a mechanism other than increased transcription of Cacna1h.

Cone synapses are enlarged with some errors in postsynaptic wiring in G369i KI mice

Earlier version: The replacement of Cav1.4 with Cav3 as the primary conduit for Ca2+ influx in G369i KI cones allowed us to dissect the importance of Cav1.4 Ca2+ signals for the molecular and structural organization of the cone synapse. To this end, we first analyzed WT and G369i KI cone synapses by immunofluorescence and confocal microscopy. Presynaptic proteins such as bassoon and members of the postsynaptic signaling complex in depolarizing (ON) cone bipolar cells (GPR179, mGluR6, and TRPM124) were enriched near cone ribbons in G369i KI mice as in WT mice (Fig. 6a). Compared to WT mice, the labeled structures in G369i KI mice occupied a larger volume (Fig. 6b,c) but were of lower intensity (Fig. 6d), suggesting an increase in the spread rather than in the levels of synaptic proteins. Unlike in WT mice, the volume of presynaptic and postsynaptic proteins increased linearly with the volume of the pedicle in G369i KI mice (Fig. 6e-i). These results show that the molecular determinants of the cone synapse can assemble in G369i KI mice but change in their sub-synaptic distribution in ways that correlate with the expansion of the cone pedicle.

Revised version: The replacement of Cav1.4 with Cav3 as the primary conduit for Ca2+ influx in G369i KI cones allowed us to dissect the importance of Cav1.4 Ca2+ signals for the molecular and structural organization of the cone synapse. To this end, we first analyzed WT and G369i KI cone synapses by immunofluorescence and confocal microscopy. Presynaptic proteins such as bassoon and members of the postsynaptic signaling complex in depolarizing (ON) cone bipolar cells (GPR179, mGluR6, and TRPM122) were enriched near cone ribbons in G369i KI mice as in WT mice (Fig. 6a). Compared to WT mice, the labeled structures in G369i KI mice occupied a larger volume (Fig. 6b,c) but were of lower intensity (Fig. 6d), suggesting an increase in the spread rather than in the levels of synaptic proteins. Unlike in WT mice, the volume of presynaptic and postsynaptic proteins increased linearly with the volume of the pedicle in G369i KI mice (Fig. 6e-i). These results show that the molecular determinants of the cone synapse can assemble in G369i KI mice but change in their sub-synaptic distribution in ways that correlate with the expansion of the cone pedicle.

Immunofluorescence characterization of cone synapses in WT and G369i KI mice.

a, Confocal images of the OPL of WT and G369i KI mice labeled with antibodies against cone arrestin (CAR) and proteins that are presynaptic (CtBP2, bassoon) or postsynaptic (GPR179, TRPM1, mGluR6). Every other panel shows high-magnification, deconvolved images of single pedicles labeled with cone arrestin (rod spherule-associated signals were removed for clarity). Arrows indicate ribbon synapses, which appear enlarged in the G369i KI pedicles. b, Violin plot representing volume of presynaptic protein labeling normalized to the volume of their respective CAR-labeled pedicles (volume occupancy). c, Violin plot representing volume of postsynaptic protein labeling normalized to the volume of their respective CAR-labeled pedicles. d, Violin plot representing fluorescence intensity of postsynaptic proteins. For b-d, **p < 0.01, ***p < 0.001, ****p < 0.0001, unpaired t-tests. e-i, Dependence of synapse size on pedicle size. Volumes corresponding to labeling of CtBP2 (e: p = 0.051, r = 0.4 for WT; p < 0.0001, r = 0.88 for G369i KI), CaV1.4 (f: p = 0.8, r = 0.06 for WT; p = 0.002, r = 0.88 for G369i KI), bassoon (g: p = 0.1, r = 0.34 for WT; p < 0.0001, r = 0.75 for G369i KI), mGluR6 (h: p = 0.32, r = 0.35 for WT; p = 0.002, r = 0.73 for G369i KI) and TRPM1 (i: p = 0.32, r = 0.35 for WT; p = 0.007, r = 0.66 for G369i KI) are plotted against pedicle volume. Dashed and solid lines represent fits by linear regression for WT and G369i KI, respectively.

Earlier version: The cone synapse is structurally complex, with the dendritic tips of two horizontal cells and an intervening ON cone bipolar cell invaginating deeply into the pedicle near the ribbon 3. To test how the switch in Cav subtypes might affect this arrangement of postsynaptic partners, we generated 3D reconstructions of cone synapses by serial block-face scanning electron microscopy (SBFSEM; Fig. 7). As with the enlarged synaptic contacts (Fig. 6), structural modifications in G369i KI pedicles were evident. Compared to WT, ribbons in G369i KI pedicles appeared disorganized and were often parallel rather than perpendicular to the presynaptic membrane (Fig.7a-c). Consistent with our confocal analyses (Fig.1), G369i KI cone pedicles extended telodendria in multiple directions rather than just apically (Fig. 7a). In addition, a slightly larger fraction of synaptic sites in the G369i KI pedicle (35% vs 14% in WT) formed incorrect postsynaptic partnerships including with glia (Fig.7b-d’; Table 2). Nevertheless, the number of ribbons is normal in G369i KI cones and about half of the ribbons make invaginating contacts with the appropriate cell types (i.e., both horizontal cells and cone bipolar cells; Table 2). On average, there were more vesicles associated with ribbons in cones of G369i KI mice (245 per ribbon) than of WT mice (178 per ribbon; Table 2), but the difference did not reach statistical significance. Together with our immunofluorescence data, these results indicate that cone synapses undergo relatively subtle morphological changes in G369i KI mice, perhaps in response to the loss of Cav1.4 Ca2+ signals.

Revised version: The cone synapse is structurally complex, with the dendritic tips of two horizontal cells and an intervening ON cone bipolar cell invaginating deeply into the pedicle near the ribbon 1. To test how the switch in Cav subtypes might affect this arrangement of postsynaptic partners, we generated 3D reconstructions of cone synapses by serial block-face scanning electron microscopy (SBFSEM; Fig. 7). As with the enlarged synaptic contacts (Fig. 6), structural modifications in G369i KI pedicles were evident. Compared to WT, ribbons in G369i KI pedicles appeared disorganized and were often parallel rather than perpendicular to the presynaptic membrane (Fig.7a-c). Consistent with our confocal analyses (Fig.1), G369i KI cone pedicles extended telodendria in multiple directions rather than just apically (Fig. 7a). In addition, a slightly larger fraction of synaptic sites in the G369i KI pedicle (35% vs 14% in WT) formed incorrect postsynaptic partnerships including with glia (Fig.7b-d’; Table 2). Nevertheless, the number of ribbons is normal in G369i KI cones and about half of the ribbons make invaginating contacts with the appropriate cell types (i.e., both horizontal cells and cone bipolar cells; Table 2). On average, there were more vesicles associated with ribbons in cones of G369i KI mice (245 per ribbon) than of WT mice (178 per ribbon; Table 2), but the difference did not reach statistical significance. Together with our immunofluorescence data, these results indicate that cone synapses undergo relatively subtle morphological changes in G369i KI mice, perhaps in response to the loss of Cav1.4 Ca2+ signals.

Three-dimensional reconstruction of cone synapses by SBFSEM.

Reconstructions were made of WT and G369i KI pedicles (n = 2 each). a, 3D renderings showing ribbons (magenta) within cone pedicles (gray) from WT and G369i KI mice. b, 3D renderings show ribbon sites in a WT cone pedicle contacting one horizontal (HC1; yellow) and three bipolar cells (BC1-3; purple). b’, Raw image from b shows a single plane example of BC1-2 and HC1 contacting the ribbon site. c-d’, Single plane raw (c,d) images and 3D reconstructions (c’,d’) show ribbon sites within the G369i KI cone pedicle contacting in c,c’: horizontal cells (HC1-2) only (upper panels), CBCs (BC1-2) only (lower panels); and in d,d’: a glial cell (orange) and an unknown partner. The glial cell completely envelops the pedicle. Inset in d’ shows glial-contacting ribbon site (arrow). In other panels, arrows indicate points of contact between ribbons and other postsynaptic elements.

Postsynaptic responses to light are present in horizontal cells of G369i KI mice but not Cav1.4 KO mice

Earlier version: Compared to Cav1.4 KO mice, the relative preservation of cone synapses in G369i KI mice allowed the unique opportunity to test whether a Cav subtype other than Cav1.4 could support synaptic release. To this end, we analyzed synaptic transmission between cones and horizontal cells (HCs). In darkness, glutamate released from cones depolarizes horizontal cells via activation of α-amino-3-hydroxy-5-methyl-4-isoxazolepropionic (AMPA)/kainate receptors25, 26. The resulting excitatory postsynaptic current (EPSC) undergoes a decline in response to light stimuli that hyperpolarize cones27. In WT HCs, a 1-s light flash (λ=410 nm) inhibited the standing EPSC, which is reflected as an outward (hyperpolarizing “ON”) current (Fig. 8a). Upon termination of the light, an inward (depolarizing “OFF”) current signaled the resumption of the EPSC (Fig. 8a). Both the ON and OFF responses in WT HCs increased with light intensity, reflecting the impact of luminance on the presynaptic membrane potential of cones and the subsequent change in glutamate release from their terminals (Fig. 8b).

Revised version: Compared to Cav1.4 KO mice, the relative preservation of cone synapses in G369i KI mice allowed the unique opportunity to test whether a Cav subtype other than Cav1.4 could support synaptic release. To this end, we analyzed synaptic transmission between cones and horizontal cells (HCs). In darkness, glutamate released from cones depolarizes horizontal cells via activation of α-amino-3-hydroxy-5-methyl-4-isoxazolepropionic (AMPA)/kainate receptors23, 24. The resulting excitatory postsynaptic current (EPSC) undergoes a decline in response to light stimuli that hyperpolarize cones25. In WT HCs, a 1-s light flash (λ=410 nm) inhibited the standing EPSC, which is reflected as an outward (hyperpolarizing “ON”) current (Fig. 8a). Upon termination of the light, an inward (depolarizing “OFF”) current signaled the resumption of the EPSC (Fig. 8a). Both the ON and OFF responses in WT HCs increased with light intensity, reflecting the impact of luminance on the presynaptic membrane potential of cones and the subsequent change in glutamate release from their terminals (Fig. 8b).

Patch clamp recordings of light responses of horizontal cells in WT, G369i KI and Cav1.4 KO mice.

Earlier version: a-b, Representative traces from whole-cell patch clamp recordings of horizontal cells held at −70 mV (a) and quantified data (b) for currents evoked by 1 s light flashes (λ = 410 nm) plotted against light intensities. In b, peak current amplitudes during (ON Response) and after (OFF Response) the light stimuli were plotted against photon flux per μm2 (Φq/μm2). Data represent mean ± SEM. WT, n = 8; G369i KI, n = 13; Cav1.4 KO, n = 8. There was a significant difference in responses of WT, G369i KI, and Cav1.4 KO horizontal cells. For ON responses, F(16, 216) = 22, p <0.0001 by two-way repeated measures ANOVA, Tukey’s multiple comparisons post-hoc test. For OFF responses, F(16, 216) = 13.61, p < 0.0001 by two-way repeated measures ANOVA, Tukey’s multiple comparisons post-hoc test. c-d, Representative traces from horizontal cells held at −70 mV (c) and quantified data (d) for currents evoked by 1-s light flashes (λ= 410 nm, 1.2 x 105 Φq/μm2) before, during, and after washout of DNQX (20 μM). In d, symbols represent responses from individual cells, n = 5 cells for WT and 3 cells for G369i KI, bars represent mean ± SEM. **, p < 0.01 by paired t-tests.

Revised version: a-b, Representative traces from whole-cell patch clamp recordings of horizontal cells held at −70 mV (a) and quantified data (b) for currents evoked by 1 s light flashes (λ = 410 nm) plotted against light intensities. In b, peak current amplitudes during (ON Response) and after (OFF Response) the light stimuli were plotted against photon flux per μm2 (Φq/μm2). Data represent mean ± SEM. WT, n = 8; G369i KI, n = 13; Cav1.4 KO, n = 8. There was a significant difference in responses of WT, G369i KI, and Cav1.4 KO horizontal cells. For ON responses, F(16, 216) = 22, p <0.0001 by two-way repeated measures ANOVA, Tukey’s multiple comparisons post-hoc test. For OFF responses, F(16, 216) = 13.61, p < 0.0001 by two-way repeated measures ANOVA, Tukey’s multiple comparisons post-hoc test. c-d, Representative traces from horizontal cells held at −70 mV (c) and quantified data (d) for currents evoked by 1-s light flashes (λ = 410 nm, 1.2 × 105 Φq/μm2) before, during, and after washout of DNQX (20 μM). In d, symbols represent responses from individual cells, n = 5 cells for WT and 3 cells for G369i KI, bars represent mean ± SEM. **, p < 0.01 by paired t-tests.

Earlier version: Consistent with the lack of mature ribbons and abnormal cone pedicles (Fig.1), HC light responses were negligible in Cav1.4 KO mice (Fig.8a,b). In contrast, the ON and OFF responses were present in G369i KI HCs although significantly lower in amplitude than in WT HCs (Fig. 8a,b). As expected, the EPSC in darkness and the light responses were abolished by the AMPA/kainate receptor antagonist, DNQX, in both WT and G369i KI HCs (Fig. 8c,d). In addition to its smaller amplitude, the transient nature of the ON response in G369i KI HCs suggested inadequate cessation of cone glutamate release by light (Fig.8b). Slow deactivation of Cav3 channels and/or their activation at negative voltages20 could give rise to Ca2+ signals that support release following light-induced hyperpolarization of G369i KI cones.

Revised version: Consistent with the lack of mature ribbons and abnormal cone pedicles (Fig.1), HC light responses were negligible in Cav1.4 KO mice (Fig.8a,b). In contrast, the ON and OFF responses were present in G369i KI HCs although significantly lower in amplitude than in WT HCs (Fig. 8a,b). As expected, the EPSC in darkness and the light responses were abolished by the AMPA/kainate receptor antagonist, DNQX, in both WT and G369i KI HCs (Fig. 8c,d). In addition to its smaller amplitude, the transient nature of the ON response in G369i KI HCs suggested inadequate cessation of cone glutamate release by light (Fig.8b). Slow deactivation of Cav3 channels and/or their activation at negative voltages18 could give rise to Ca2+ signals that support release following light-induced hyperpolarization of G369i KI cones.

Earlier version: To test if the diminished HC light responses correlated with lower presynaptic Ca2+ signals in G369i KI cones, we performed 2-photon imaging of cone pedicles filled with the Ca2+ indicator Fluo3 and labeled with Alexa-568-conjugated peanut agglutinin. In WT cones, depolarization by K+ (50 mM) caused a sustained increase of Ca2+ (ΔF/F0 = 1.85 ± 0.1, n = 11 terminals) that was significantly lower after the addition of isradipine (ΔF/F0 = 1.44 ± 0.07, p = 0.0015 by repeated measures one-way ANOVA with Tukey’s multiple comparison posthoc test; F[1.608, 16.08] = 55.17; p < 0.0001; Fig.9a,b). The Ca2+ signal remaining in the presence of isradipine could arise from unblocked Cav1.4 channels and/or Ca2+ release from intracellular stores28. In G369i KI cones, the peak amplitude of the K+-evoked Ca2+ signal (ΔF/F0 = 1.07 ± 0.02, n = 14 terminals) was significantly lower than in WT cones (t(23)=8.74, p < 0.0001 by unpaired t-test, n = 14 terminals). The decrease in the Ca2+ signal in the G369i KI cone pedicles during isradipine exposure paralleled that during the depolarization (Fig.9a), suggesting it was not a consequence of Cav1 blockade. The lower amplitude and more transient Ca2+ signals mediated by Cav3 channels may be cleared more rapidly in G369i KI cones as compared to prolonged Cav1.4-initiated Ca2+ signals in WT cones. Taken together, our results suggest that Cav3 channels nominally support Ca2+ signals and synaptic transmission in cones of G369i KI mice.

Revised version: To test if the diminished HC light responses correlated with lower presynaptic Ca2+ signals in G369i KI cones, we performed 2-photon imaging of vertical slices prepared from whole retina that was incubated with the Ca2+ indicator Fluo3-AM and Alexa-568-conjugated peanut agglutinin (PNA) to demarcate regions of interest (ROIs) corresponding to cone pedicles. With this approach Fluo3 fluorescence was detected only in photoreceptors and ganglion cells and not inner retinal cell-types (e.g., horizontal cells, bipolar cells, Mueller cell soma). Thus, Ca2+ signals reported by Fluo3 fluorescence near PNA-labeling originated primarily from cones. In WT cones, depolarization by K+ (50 mM) caused a sustained increase of Ca2+ (ΔF/F0 = 1.85 ± 0.1, n = 11 terminals) that was significantly lower after the addition of isradipine (ΔF/F0 = 1.44 ± 0.07, p = 0.0015 by repeated measures one-way ANOVA with Tukey’s multiple comparison posthoc test; F[1.608, 16.08] = 55.17; p < 0.0001; Fig.9a,b). The Ca2+ signal remaining in the presence of isradipine could arise from unblocked Cav1.4 channels and/or Ca2+ release from intracellular stores26. In G369i KI cones, the peak amplitude of the K+-evoked Ca2+ signal (ΔF/F0 = 1.07 ± 0.02, n = 14 terminals) was significantly lower than in WT cones (t(23)=8.74, p < 0.0001 by unpaired t-test, n = 14 terminals). The decrease in the Ca2+ signal in the G369i KI cone pedicles during isradipine exposure paralleled that during the depolarization (Fig.9a), suggesting it was not a consequence of Cav1 blockade. The lower amplitude and more transient Ca2+ signals mediated by Cav3 channels may be cleared more rapidly in G369i KI cones as compared to prolonged Cav1.4-initiated Ca2+ signals in WT cones. Taken together, our results suggest that Cav3 channels nominally support Ca2+ signals and synaptic transmission in cones of G369i KI mice.

Added in the revised version: Ca2+ imaging of cone terminals in WT and G369i KI mice.

Added in the revised version: Depolarization-evoked fluorescence changes (ΔF) were measured by two-photon imaging of Fluo3-filled cone pedicles labeled with PNA-Alexa 568. a, Representative images at the indicated timepoints (T) before (control), during, and after (wash) bath application of KCl (50 mM) and KCl+ISR (10 μM). For clarity, PNA-Alexa 568 labeling was outlined (magenta, solid circles) and the actual labeling only shown in the initial panel of images. Dashed, yellow circles represent ROIs for analyses. b, Changes in cone pedicle Fluo3 fluorescence were normalized to the initial fluorescence intensity at time= 0 (ΔF/F0). Solid lines and shaded areas represent the mean ± SEM, respectively. WT, n = 11 pedicles; G369i KI, n = 14 pedicles.

Light responses of bipolar cells and visual behavior are present in G369i KI but not Cav1.4 KO mice

Earlier version: While HCs play important roles in lateral inhibition of photoreceptor output, the vertical dissemination of visual information from cones to the inner retina is relayed by glutamate to ON and OFF cone bipolar cells (CBCs) which express mGluR6 and AMPA/kainate receptors, respectively. To test for the existence of cone to CBC transmission in G369i KI mice, we recorded electroretinograms (ERGs) under light-adapted conditions using Cav1.4 KO mice as a negative control (Fig. 10a-d). In these recordings, the light-induced response of photoreceptors and the postsynaptic response of ON CBCs corresponds to the a- and b-waves, respectively. Unlike in Cav1.4 KO mice, the a-waves of G369i KI mice resembled those in WT mice, which indicates that cones do not degenerate in this mouse strain (Fig. 10a,b). While greatly reduced in amplitude, the b-wave was measurable in G369i KI mice and significantly larger than in Cav1.4 KO mice at the highest light intensities (F(14,104)=41, p<0.0001, post-hoc Tukey’s p=0.0038 at 2.3 log cd·s/m2; Fig. 10a,b). We also recorded flicker ERGs using 10 Hz light stimuli that can isolate cone pathways involving both ON and OFF CBCs29, 30. In WT mice, flicker responses exhibited two peaks, one at a lower irradiance (−2 log cd·s/m2) and one at higher irradiance (0.5 log cd·s/m2; Fig.10c,d). The peak at the lower irradiance is attributed to responses in both rod and cone pathways, and the peak at higher irradiance is attributed to responses exclusively in cone pathways29, 30. Flicker responses in G369i KI mice showed a similar non-monotonic relation as in WT mice but were significantly lower in amplitude (F(26,143)=51.18, p<0.0001, post-hoc Tukey’s p=0.0237 at 2.5 log cd·s/m2; Fig.10c,d). The inverted flicker responses at higher illuminations in G369i KI mice were absent in Cav1.4 KO mice and may result from the hyperpolarizing contribution of cone-to-OFF CBC transmission (Fig.10c).

Revised version: While HCs play important roles in lateral inhibition of photoreceptor output, the vertical dissemination of visual information from cones to the inner retina is relayed by glutamate to ON and OFF cone bipolar cells (CBCs) which express mGluR6 and AMPA/kainate receptors, respectively. To test for the existence of cone to CBC transmission in G369i KI mice, we recorded electroretinograms (ERGs) under light-adapted conditions using Cav1.4 KO mice as a negative control (Fig. 10a-d). In these recordings, the light-induced response of photoreceptors and the postsynaptic response of ON CBCs corresponds to the a- and b-waves, respectively. Unlike in Cav1.4 KO mice, the a-waves of G369i KI mice resembled those in WT mice, which indicates that cones do not degenerate in this mouse strain (Fig. 10a,b). While greatly reduced in amplitude, the b-wave was measurable in G369i KI mice and significantly larger than in Cav1.4 KO mice at the highest light intensities (F(14,104)=41, p<0.0001, post-hoc Tukey’s p=0.0038 at 2.3 log cd·s/m2; Fig. 10a,b). We also recorded flicker ERGs using 10 Hz light stimuli that can isolate cone pathways involving both ON and OFF CBCs27, 28. In WT mice, flicker responses exhibited two peaks, one at a lower irradiance (−2 log cd·s/m2) and one at higher irradiance (0.5 log cd·s/m2; Fig.10c,d). The peak at the lower irradiance is attributed to responses in both rod and cone pathways, and the peak at higher irradiance is attributed to responses exclusively in cone pathways27, 28. Flicker responses in G369i KI mice showed a similar non-monotonic relation as in WT mice but were significantly lower in amplitude (F(26,143)=51.18, p<0.0001, post-hoc Tukey’s p=0.0237 at 2.5 log cd·s/m2; Fig.10c,d). The inverted flicker responses at higher illuminations in G369i KI mice were absent in Cav1.4 KO mice and may result from the hyperpolarizing contribution of cone-to-OFF CBC transmission (Fig.10c).

Removed in the revised version: To test whether the cone synaptic responses in G369i KI mice enable vision-guided behavior, we used a swim test that assesses the ability of mice to identify a visible platform31. Under scotopic (dark) conditions, G369i KI and Cav1.4 KO mice took significantly longer to find the platform compared to WT mice (Fig. 10e,f). These results are consistent with flicker response assays (Fig.10c,d) as well as the absence of ICa in rods and rod-to-rod bipolar cell synaptic transmission in G369i KI mice19. Under photopic (light) conditions, G369i KI mice but not Cav1.4 KO mice performed as well as WT mice (Fig. 10e,f). Thus, G369i KI mice have enough visual function under photopic conditions to quickly find the platform. Collectively, our results suggest that Cav3 channels can support cone synaptic responses that are sufficient for visual behavior in G369i KI but not Cav1.4 KO mice.

Removed in the revised version: Discussion

Removed in the revised version: The nervous system has remarkable abilities to adapt to pathological perturbations in neuronal activity. In the retina, ablation or degeneration of photoreceptors triggers various postsynaptic mechanisms that maintain some level of visual function in rod or cone pathways. These include remodeling of bipolar cell dendrites and their synapses32–34 as well as changes in the sensitivity of bipolar cells to photoreceptor input and inhibitory modulation34, 35. To our knowledge, this study provides the first evidence for a presynaptic form of homeostatic plasticity that originates within photoreceptors. Using Cav1.4 KO and G369i KI mice, we identify the upregulation of a Cav3 conductance as a common response to Cav1.4 loss-of-function in cones. However, Cav3 channels can only compensate for Cav1.4 loss-of-function if the basic aspects of cone synapse structure are maintained through expression of the G369i mutant channel. Thus, our results also highlight a crucial, non-conducting role for the Cav1.4 protein that allows cone synapses to function in the absence of Cav1.4 Ca2+ signals.

Removed in the revised version: A non-canonical role for Cav1.4 in regulating cone synapse assembly

Removed in the revised version: As shown for rod synapses19, ribbons and other components of the pre- and post-synaptic complex assemble at cone synapses in G369i KI mice (Figs. 1,6,7). Cav3 Ca2+ signals are dispensable for this process since their presence in Cav1.4 KO cones is not accompanied by any semblance of ribbon synapses (Fig.1b). An intimate relationship between Cav1.4 and ribbons is supported by the colocalization of Cav1.4 and puncta resembling ribbon precursor spheres in the developing OPL7. Cav1.4 interacts directly or indirectly with a variety of ribbon synapse associated proteins including RIM 36, CAST/ERC2 37, bassoon38, unc11939–41, and a major component of the ribbon, RIBEYE42. Cav1.4 could pioneer sites of ribbon assembly perhaps by serving as a docking or nucleation site for the active zone. At the same time, ribbon-associated proteins help cluster Cav1.4 near ribbons since knockout of these proteins in mice leads to fewer presynaptic Cav1.4 channels and in some cases shortened ribbons43. Regardless of the mechanism, our findings show that the formation of photoreceptor synaptic complexes does not require Ca2+ influx through Cav1.4 channels.

Removed in the revised version: Functional Cav3 channels are present in cones in the absence of Cav1.4 Ca2+ signals

Removed in the revised version: Our drop-seq analysis aligns with previous studies indicating that Cacna1h is expressed at extremely low levels compared to Cacna1f in WT mouse cones (Supp.Fig.S3)23, 44. However, in contrast to recordings of WT mouse cone pedicles in a previous study21, we found no evidence for Cav3-mediated currents in somatic recordings of cones in WT mice (Figs.2,3). We do not feel the discrepancy results from the different recording configurations given the electrically compact nature of cones. Moreover, recordings from macaque cone pedicles and from the ellipsoids of cylindrically-shaped ground squirrel cones also did not reveal a Cav3-mediated current (Figs.4,5). We acknowledge that there are caveats of the pharmacological approach used in our study as well as the previous work21. For example, dihydropyridine antagonists such as isradipine have relatively low affinity for some Cav1 subtypes, including Cav1.422, 45 and have strongly voltage-dependent blocking properties. At a holding voltage of −90 mV, isradipine at 1 μM causes only ∼80% inhibition of Cav1.4 with greater block at depolarized voltages22, which agrees with our recordings of cones in WT mice (Fig.3a) and ground squirrels (Fig.4b-e). Moreover, in the low micromolar range, Cav3 blockers such as Z944 46 and ML 218 (Supp.Fig. S2) have nonspecific effects on Cav1 current amplitude and activation voltage. Thus, Cav channels mediating a ICa that is suppressed by Cav3 blockers and spared by dihydropyridines at negative voltages may be mistakenly categorized as Cav3 subtypes. In our experiments, the biophysical properties of the isradipine-sensitive ICa in cones of WT mice and ground squirrels resembled only those of Cav1, whereas those of the ML 218-sensitive ICa in cones of G369i KI mice resembled only those of Cav3. Therefore, we favor the interpretation that Cav3 contributes to ICa in mouse cones only upon silencing of Cav1.4 either by KO of the corresponding protein or by silencing its conductance.

Removed in the revised version: Since Cacna1h expression was similar in WT and G369i KI cones (Supp.Fig.S3), what causes the increased Cav3 channel activity in the latter? Among the Cav3 subtypes, Cav3.2 is uniquely sensitive to post-translational modifications that facilitate the opening of the channel and/or its cell surface trafficking47. For example, Cav3.2 interacts with the deubiquitinating enzyme USP5 which enhances levels of Cav3.2 protein levels by suppressing its proteasomal degradation48. In addition, Cav3.2 is phosphorylated by calmodulin-dependent kinase II49 and cyclin-dependent kinase 550, which enhances channel activation and current density, respectively. Further interrogation of our drop-seq dataset could illuminate novel Cav3.2-regulatory pathways that are modified in G369i KI cones.

Removed in the revised version: Consequences of Cav1.4 loss of function and relevance for CSNB2

Removed in the revised version: Despite having the same complement of proteins as in WT mice, cone synapses in G369i KI mice were enlarged and made errors in postsynaptic partner selection (Figs.6,7). While I-V curves indicate similar peak ICa amplitudes in cones of WT and G369i KI mice near the membrane potential of cones in darkness (−45 to −50 mV51, Fig.2a,c), the strong inactivation of Cav3 channels and/or reduced Ca2+ release from intracellular stores likely account for the greatly diminished Ca2+ signals in G369i KI pedicles (Fig.9). Paradoxically, reductions in presynaptic Ca2+ in rod photoreceptors correlated with illumination-dependent shrinkage of ribbons in vivo and in vitro2, 52, 53. However, ribbon synapse size increases with low presynaptic Ca2+ in zebrafish sensory hair cells54. The mechanism involves decreased mitochondrial Ca2+ uptake and increase in the redox state of nicotinamide adenine dinucleotide (NAD+/NADH ratio) 55. Cav3 Ca2+ signals may decay too quickly in G369i KI cones (Fig.9) to enable mitochondrial Ca2+ uptake or other mechanisms that trim ribbons.

Removed in the revised version: Ca2+ imaging of cone terminals in WT and G369i KI mice.

Removed in the revised version: Depolarization-evoked fluorescence changes (ΔF) were measured by two-photon imaging of Fluo3-filled cone pedicles labeled with PNA-Alexa 568.

Removed in the revised version: a, Representative images at the indicated timepoints (T) before (control), during, and after (wash) bath application of KCl (50 mM) and KCl+ISR (10 μM). For clarity, PNA-Alexa 568 labeling was outlined (magenta, solid circles) and the actual labeling only shown in the initial panel of images. Dashed, yellow circles represent ROIs for analyses. b, Changes in cone pedicle Fluo3 fluorescence were normalized to the initial fluorescence intensity at time= 0 (ΔF/F0). Solid lines and shaded areas represent the mean ± SEM, respectively. WT, n = 11 pedicles; G369i KI, n = 14 pedicles.

Removed in the revised version: An intriguing result was the linear increase in ribbon synapse size with pedicle volume in G369i KI but not in WT mice (Fig. 6e-i), which could represent a form of homeostatic synaptic scaling as has been demonstrated at the neuromuscular junction56, 57. Similarly, the presence of some non-invaginating contacts with incorrect partner pairings in G369i KI mice (Fig.7, Table 2) could be an attempt to increase synaptic connectivity to compensate for weakened synaptic output. Consistent with this possibility, diminished presynaptic Ca2+ signals and non-invaginating cone synapses also characterize the retina of mice lacking the extracellular α2δ−4 subunit of Cav1.458. The cone synapses of G369i KI mice and α2δ−4 KO mice are clearly functional given that photopic visual behavior is preserved in both mouse strains (Fig.10e,f)58 and, in the case of G369i KI mice, HCs exhibit EPSCs whose amplitude decreases with light and AMPA/kainate receptor blockade (Fig.8). Defects in cone synaptic transmission that correlate with aberrant cone synapse morphology and connectivity also typify a mouse strain with a gain-of function mutation in Cav1.4 linked to CSNB259, 60. Thus, balanced presynaptic Ca2+ signals may be needed to encode the correct post-synaptic wiring and structure of the cone synapse.

Removed in the revised version: In addition to structural abnormalities in cone synapses, the inability of Cav3 channels to support sustained presynaptic Ca2+ entry likely contributed to the impaired light responses in HCs and CBCs in G369i KI mice (Figs.8,10). Due to their slow activation and strong inactivation 61, Cav3 channels are not expected to fuel the Ca2+ nanodomains that support fast and sustained components of release, both of which occur only at ribbon sites in cones1 62. Cav3 channels may also be located further from the ribbon than Cav1.4, thus lowering the efficiency of coupling to exocytosis. Regardless of the features that differentiate the contributions of Cav3 and Cav1.4 to cone synaptic release, our findings provide further support for the widely held view that Cav1 channels are well-suited to support transmission at ribbon-type synapses.

Removed in the revised version: Based on extremely heterogeneous clinical presentations, CSNB2 manifests as a spectrum of visual disorders that originate from various mutations in CACNA1F 14, 63. Even though some CSNB2 mutant Cav1.4 channels may traffic normally to the plasma membrane, many of these mutations are expected to produce non-functional, non-conducting Cav1.4 channels64, 65. Yet, the visual phenotypes of CSNB2 patients are not as severe as the complete blindness in Cav1.4 KO mice, which lack any Cav1.4 protein expression and exhibit no signs of visual behavior (Fig.10e,f)58. Collectively, our results suggest that G369i KI mice accurately model Cav1.4 channelopathies in CSNB2 patients that are characterized by a greater impairment in rod than in cone pathways. This interpretation is supported by our findings that G369i KI mice exhibit horizontal cell responses to bright but not dim illumination (Fig. 8), ERG responses under conditions of light adaptation (Fig.10a-d) but not dark-adaptation 19, and visual behavior under photopic but not scotopic conditions (Fig. 10e,f). Together with the enlargement of synaptic sites, modest levels of synaptic release from cones of G369i KI mice may support nominal transmission of visual information through cone pathways. G369i KI mice could also exhibit homeostatic alterations in the inner retina, which are known to support visual function when photoreceptor input is severely compromised 33–35. Future studies of the retinal circuitry and visual behavior of G369i KI mice could identify compensatory pathways that are recruited upon Cav1.4 loss-of-function and how they might be targeted in novel therapies for CSNB2 and related disorders.

Characterization of ERGs and visual behavior in WT, G369i KI, and Cav1.4 KO mice.

a, Representative traces of photopic ERGs recorded in the presence of background green light (20 cd · s/m2) in WT, G369i KI and CaV1.4 KO mice. Flash intensities are shown at left. Arrows and arrowheads depict the a- and b-waves, respectively. b, a-wave (left) and b-wave (right) amplitudes are plotted against light intensity. Symbols and bars represent mean ± SEM. WT, n = 7; G369i KI, n = 5; CaV1.4 KO, n = 6. *, p < 0.05; **, p < 0.01; ***, p < 0.001; ****, p < 0.0001; ns, not significant; Two-way ANOVA with Tukey’s posthoc multiple comparisons. c,d, Representative traces (c) and quantified data (d) for 10 Hz flicker responses evoked by white light flashes of increasing luminance (from −4 to 2 log cd· s/m2). Arrows in c depict inverted waveform responses in G369i KI mice that are absent in Cav1.4 KO mice. Symbols and bars represent mean ± SEM. WT, n = 7; G369i KI, n = 5; CaV1.4 KO, n = 6. *, p < 0.05; **, p < 0.01. Two-way ANOVA with Tukey’s posthoc multiple comparisons. e, Representative swim path traces of WT, G369i KI and CaV1.4 KO mice from the visible platform swim tests performed in the dark (scotopic, upper panel) and light (photopic, lower panel). f, Time required to reach the platform (latency to platform) was compared for each mouse strain. Symbols represent the average of the last 3 swim trials for each mouse of each genotype for both dark and light conditions. Dotted lines represent the mean. WT, n = 11; G369i KI, n = 10; CaV1.4 KO, n = 9. *, p < 0.05; ***, p < 0.001; ****, p < 0.0001; Kruskal-Wallis one-way ANOVA with Dunn’s post hoc analysis.

Added in the revised version: To test whether the cone synaptic responses in G369i KI mice enable vision-guided behavior, we used a swim test that assesses the ability of mice to identify a visible platform29. Under scotopic (dark) conditions, G369i KI and Cav1.4 KO mice took significantly longer to find the platform compared to WT mice (Fig. 10e,f). These results are consistent with flicker response assays (Fig.10c,d) as well as the absence of ICa in rods and rod-to-rod bipolar cell synaptic transmission in G369i KI mice17. Under photopic (light) conditions, G369i KI mice but not Cav1.4 KO mice performed as well as WT mice (Fig. 10e,f). Thus, G369i KI mice have enough visual function under photopic conditions to quickly find the platform. Collectively, our results suggest that Cav3 channels can support cone synaptic responses that are sufficient for visual behavior in G369i KI but not Cav1.4 KO mice.

Added in the revised version: Discussion

Added in the revised version: The nervous system has remarkable abilities to adapt to pathological perturbations in neuronal activity. In the retina, ablation or degeneration of photoreceptors triggers various postsynaptic mechanisms that maintain some level of visual function in rod or cone pathways. These include remodeling of bipolar cell dendrites and their synapses30–32 as well as changes in the sensitivity of bipolar cells to photoreceptor input and inhibitory modulation32, 33. To our knowledge, this study provides the first evidence for a presynaptic form of homeostatic plasticity that originates within photoreceptors. Using Cav1.4 KO and G369i KI mice, we identify the upregulation of a Cav3 conductance as a common response to Cav1.4 loss-of-function in cones. However, Cav3 channels can only compensate for Cav1.4 loss-of-function if the basic aspects of cone synapse structure are maintained through expression of the G369i mutant channel. Thus, our results also highlight a crucial, non-conducting role for the Cav1.4 protein that allows cone synapses to function in the absence of Cav1.4 Ca2+ signals.

Added in the revised version: A non-canonical role for Cav1.4 in regulating cone synapse assembly

Added in the revised version: As shown for rod synapses17, ribbons and other components of the pre- and post-synaptic complex assemble at cone synapses in G369i KI mice (Figs. 1,6,7). Cav3 Ca2+ signals are dispensable for this process since their presence in Cav1.4 KO cones is not accompanied by any semblance of ribbon synapses (Fig.1b). An intimate relationship between Cav1.4 and ribbons is supported by the colocalization of Cav1.4 and puncta resembling ribbon precursor spheres in the developing OPL5. Cav1.4 interacts directly or indirectly with a variety of ribbon synapse associated proteins including RIM 34, CAST/ERC2 35, bassoon36, unc11937–39, and a major component of the ribbon, RIBEYE40. Cav1.4 could pioneer sites of ribbon assembly perhaps by serving as a docking or nucleation site for the active zone. At the same time, ribbon-associated proteins help cluster Cav1.4 near ribbons since knockout of these proteins in mice leads to fewer presynaptic Cav1.4 channels and in some cases shortened ribbons41. Regardless of the mechanism, our findings show that the formation of photoreceptor synaptic complexes does not require Ca2+ influx through Cav1.4 channels.

Added in the revised version: Functional Cav3 channels are present in cones in the absence of Cav1.4 Ca2+ signals

Added in the revised version: Our drop-seq analysis aligns with previous studies indicating that Cacna1h is expressed at extremely low levels compared to Cacna1f in WT mouse cones (Supp.Fig.S3)21, 42. However, in contrast to recordings of WT mouse cone pedicles in a previous study19, we found no evidence for Cav3-mediated currents in somatic recordings of cones in WT mice (Figs.2,3). We do not feel the discrepancy results from the different recording configurations given the electrically compact nature of cones. Moreover, recordings from macaque cone pedicles and from the ellipsoids of cylindrically-shaped ground squirrel cones also did not reveal a Cav3-mediated current (Figs.4,5). We acknowledge that there are caveats of the pharmacological approach used in our study as well as the previous work19. For example, dihydropyridine antagonists such as isradipine have relatively low affinity for some Cav1 subtypes, including Cav1.420, 43 and have strongly voltage-dependent blocking properties. At a holding voltage of −90 mV, isradipine at 1 μM causes only ∼80% inhibition of Cav1.4 with greater block at depolarized voltages20, which agrees with our recordings of cones in WT mice (Fig.3a) and ground squirrels (Fig.4b-e). Moreover, in the low micromolar range, Cav3 blockers such as Z944 44 and ML 218 (Supp.Fig. S2) have nonspecific effects on Cav1 current amplitude and activation voltage. Thus, Cav channels mediating a ICa that is suppressed by Cav3 blockers and spared by dihydropyridines at negative voltages may be mistakenly categorized as Cav3 subtypes. In our experiments, the biophysical properties of the isradipine-sensitive ICain cones of WT mice and ground squirrels resembled only those of Cav1, whereas those of the ML 218-sensitive ICa in cones of G369i KI mice resembled only those of Cav3. Therefore, we favor the interpretation that Cav3 contributes to ICa in mouse cones only upon silencing of Cav1.4 either by KO of the corresponding protein or by silencing its conductance.

Added in the revised version: Since Cacna1h expression was similar in WT and G369i KI cones (Supp.Fig.S3), what causes the increased Cav3 channel activity in the latter? Among the Cav3 subtypes, Cav3.2 is uniquely sensitive to post-translational modifications that facilitate the opening of the channel and/or its cell surface trafficking45. For example, Cav3.2 interacts with the deubiquitinating enzyme USP5 which enhances levels of Cav3.2 protein levels by suppressing its proteasomal degradation46. In addition, Cav3.2 is phosphorylated by calmodulin-dependent kinase II47 and cyclin-dependent kinase 548, which enhances channel activation and current density, respectively. Further interrogation of our drop-seq dataset could illuminate novel Cav3.2-regulatory pathways that are modified in G369i KI cones.

Added in the revised version: Consequences of Cav1.4 loss of function and relevance for CSNB2

Added in the revised version: Despite having the same complement of proteins as in WT mice, cone synapses in G369i KI mice were enlarged and made errors in postsynaptic partner selection (Figs.6,7). While I-V curves indicate similar peak ICa amplitudes in cones of WT and G369i KI mice near the membrane potential of cones in darkness (−45 to −50 mV49, Fig.2a,c), the strong inactivation of Cav3 channels and/or reduced Ca2+ release from intracellular stores likely account for the greatly diminished Ca2+ signals in G369i KI pedicles (Fig.9). Paradoxically, reductions in presynaptic Ca2+ in rod photoreceptors correlated with illumination-dependent shrinkage of ribbons in vivo and in vitro50–52. However, ribbon synapse size increases with low presynaptic Ca2+ in zebrafish sensory hair cells53. The mechanism involves decreased mitochondrial Ca2+ uptake and increase in the redox state of nicotinamide adenine dinucleotide (NAD+/NADH ratio)54. Cav3 Ca2+ signals may decay too quickly in G369i KI cones (Fig.9) to enable mitochondrial Ca2+ uptake or other mechanisms that trim ribbons.

Added in the revised version: An intriguing result was the linear increase in ribbon synapse size with pedicle volume in G369i KI but not in WT mice (Fig. 6e-i), which could represent a form of homeostatic synaptic scaling as has been demonstrated at the neuromuscular junction55, 56. Similarly, the presence of some non-invaginating contacts with incorrect partner pairings in G369i KI mice (Fig.7, Table 2) could be an attempt to increase synaptic connectivity to compensate for weakened synaptic output. Consistent with this possibility, diminished presynaptic Ca2+ signals and non-invaginating cone synapses also characterize the retina of mice lacking the extracellular α2δ−4 subunit of Cav1.457. The cone synapses of G369i KI mice and α2δ−4 KO mice are clearly functional given that photopic visual behavior is preserved in both mouse strains (Fig.10e,f)57 and, in the case of G369i KI mice, HCs exhibit EPSCs whose amplitude decreases with light and AMPA/kainate receptor blockade (Fig.8). Defects in cone synaptic transmission that correlate with aberrant cone synapse morphology and connectivity also typify a mouse strain with a gain-of function mutation in Cav1.4 linked to CSNB258, 59. Thus, balanced presynaptic Ca2+ signals may be needed to encode the correct post-synaptic wiring and structure of the cone synapse.

Added in the revised version: In addition to structural abnormalities in cone synapses, the inability of Cav3 channels to support sustained presynaptic Ca2+ entry likely contributed to the impaired light responses in HCs and CBCs in G369i KI mice (Figs.8,10). Due to their slow activation and strong inactivation 60, Cav3 channels are not expected to fuel the Ca2+ nanodomains that support fast and sustained components of release, both of which occur only at ribbon sites in cones61 62. Cav3 channels may also be located further from the ribbon than Cav1.4, thus lowering the efficiency of coupling to exocytosis. Regardless of the features that differentiate the contributions of Cav3 and Cav1.4 to cone synaptic release, our findings provide further support for the widely held view that Cav1 channels are well-suited to support transmission at ribbon-type synapses.

Added in the revised version: Based on extremely heterogeneous clinical presentations, CSNB2 manifests as a spectrum of visual disorders that originate from various mutations in CACNA1F 12, 63. Even though some CSNB2 mutant Cav1.4 channels may traffic normally to the plasma membrane, many of these mutations are expected to produce non-functional, non-conducting Cav1.4 channels64, 65. Yet, the visual phenotypes of CSNB2 patients are not as severe as the complete blindness in Cav1.4 KO mice, which lack any Cav1.4 protein expression and exhibit no signs of visual behavior (Fig.10e,f)57. Collectively, our results suggest that G369i KI mice accurately model Cav1.4 channelopathies in CSNB2 patients that are characterized by a greater impairment in rod than in cone pathways. This interpretation is supported by our findings that G369i KI mice exhibit horizontal cell responses to bright but not dim illumination (Fig. 8), ERG responses under conditions of light adaptation (Fig.10a-d) but not dark-adaptation 17, and visual behavior under photopic but not scotopic conditions (Fig. 10e,f). Together with the enlargement of synaptic sites, modest levels of synaptic release from cones of G369i KI mice may support nominal transmission of visual information through cone pathways. G369i KI mice could also exhibit homeostatic alterations in the inner retina, which are known to support visual function when photoreceptor input is severely compromised 31–33. Future studies of the retinal circuitry and visual behavior of G369i KI mice could identify compensatory pathways that are recruited upon Cav1.4 loss-of-function and how they might be targeted in novel therapies for CSNB2 and related disorders.

Added in the revised version: Supporting information